---
title: "Exercise 10"
output: learnr::tutorial
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
list.of.packages <- c("learnr","shiny","tidyverse","lme4","lmerTest")
new.packages <- list.of.packages[!(list.of.packages %in% installed.packages()[,"Package"])]
if(length(new.packages)) install.packages(new.packages)

library(learnr)
require(tidyverse)
require(lme4)
require(lmerTest)
require(remotes)
if(!("gradethis" %in% installed.packages())){remotes::install_github("rstudio/gradethis")}
require(gradethis)

knitr::opts_chunk$set(echo = FALSE)

gradethis::gradethis_setup()

threat.data <- data.frame(participant_id = 1:40,
	locus_of_control = c(0.67,0.72,-1.67,-0.15,0.43,1.45,-1.49,-1.44,-0.01,1.02,-0.66,0.06,-0.13,-0.17,0.35,-1.02,0.49,-0.45,0.09,0.79,0.01,-0.53,0.18,0.52,0.14,-2.09,0.82,0.14,-1.84,-0.21,0.54,1.15,2.66,0.08,-0.92,-0.13,-0.52,1.24,0.69,1.77),
	face_1 = c(0.1,-0.73,1.14,-0.46,-0.72,-0.72,-0.16,0.77,0.23,-1.16,1.11,-0.58,0.73,0.1,0.17,0.89,-1.01,1.22,0.07,-0.18,-0.36,0.72,-0.16,-0.91,-0.57,0.97,-0.05,-0.57,0.76,-0.05,-0.36,-0.39,-1.16,-0.09,-0.56,-0.01,-0.23,-1.07,-1.26,-0.47),
	face_2 = c(1.4,0.04,0.14,0.41,-0.61,0.43,-0.87,-0.21,-0.04,0.07,-0.28,-0.25,-0.58,-0.22,0.85,0.24,-0.37,0.22,0.07,-0.05,0.51,-0.93,0.32,0.57,0.62,-0.36,0.54,0.09,-1.24,0.22,0.37,0.79,0.45,0.74,-0.13,-0.02,-0.83,0.38,0.17,1.42),
	face_3 = c(-0.47,-0.33,-0.46,-0.6,0.04,0.31,0.1,1.21,-0.46,0.04,-0.28,0.78,-0.84,-0.19,-0.05,-0.07,0.88,-0.47,-0.78,-0.18,0.64,0.42,0.09,0.13,-0.09,-0.14,0.15,0.07,-1,-0.34,0.56,-0.6,1.11,0.56,0.17,0.38,-0.65,1.18,0.63,-0.07),
	face_4 = c(0.17,0.08,-0.42,0.09,-0.15,-0.41,-0.15,0.35,-0.51,0.19,-0.11,0.47,1.59,-0.24,-0.97,-0.3,0.11,-1.01,-0.49,-0.48,-0.24,-0.83,-0.26,-0.05,-0.37,-0.55,0.56,-0.42,-0.14,0.13,-0.31,1.29,0.2,-0.21,0.4,-0.31,0,0.4,-0.21,-0.21),
	face_5 = c(-0.76,-0.33,0.47,-0.69,-0.6,0.62,0.1,-0.12,-0.4,0.54,-0.1,1.51,0.61,0.02,0.31,0.17,0.39,-0.29,0.19,-0.26,0.53,-1.02,0.17,0.19,-0.36,-0.71,-0.42,0.27,-0.49,0.22,0.32,0.9,0.44,-0.68,0.01,-0.37,-0.81,0.25,0.05,0.04),
	face_6 = c(-0.55,-0.58,0.67,0.1,-0.02,0.06,0.33,0.18,0.68,-0.54,-0.16,-1.74,-1.05,0.29,0.16,-0.2,-0.47,0.14,0.51,-0.3,0.17,0.56,-0.51,-0.41,-0.03,0.01,-0.22,0.44,0.33,-0.18,0.29,0.32,0.38,0.15,-0.25,0.27,0.58,0.51,-0.68,0.5),
	face_7 = c(-0.74,0.18,0.46,0.48,-0.78,-0.5,0,0.32,0.32,-0.36,0.24,-1.1,1.01,0.57,-0.36,0.73,-0.03,0.62,-0.56,0.51,-1.25,-0.33,-0.33,-0.34,-0.24,0.92,-1.02,0.15,-0.97,0.33,0.28,-0.39,-1.21,0.21,0.19,0.08,0.42,-0.65,-0.29,-0.74),
	face_8 = c(-0.75,-1.25,1.84,-0.12,-0.52,-1.57,0.91,0.3,0.32,-1.19,0.15,0.81,0.01,0.59,-1.18,0.97,-0.99,0.35,-0.04,-1.13,-1.04,0.83,0.35,-0.97,-0.19,1.78,-0.87,-0.69,0.38,-0.64,-0.45,-1.86,-1.41,-1.26,0.84,0.06,-0.36,-0.77,-0.31,-1.63),
	face_9 = c(0.65,0.31,-1.57,-0.3,-0.33,0.47,-0.97,-1.16,0.41,0.74,0.12,0.47,0.22,0.57,0.33,-0.86,0.58,0.06,-0.27,0.55,0.51,-0.1,-0.05,0.17,-0.44,-1.14,1.16,0.26,-1.38,1.24,1.32,0.75,0.98,-0.01,-0.39,-0.59,-0.44,0.41,-0.61,0.96),
	face_10 = c(-0.24,-0.49,0.52,-0.16,0.32,-0.95,0.72,0.48,-0.24,-0.82,0.27,0.14,-0.09,-0.89,-0.54,0.09,-0.91,0.02,-0.77,-0.51,-0.89,0.41,0.51,-0.48,-0.1,-0.03,-0.29,-0.05,1.26,0.06,-0.73,-0.12,-1.82,-0.48,0.57,-0.52,0.55,-0.1,-1.08,-1.67),
	face_11 = c(0.43,-0.53,0.1,-0.3,0.81,-1.6,-0.81,0.97,0.49,-0.81,-0.19,0.15,0.22,0.56,-0.52,-0.35,0.05,-0.62,-0.32,0.16,0.26,-0.13,-0.04,-0.7,0.93,0.82,-0.39,-0.34,0.54,0.01,0.21,-0.6,-1.23,0.07,-0.73,0.05,0.09,0.47,-0.19,-0.28),
	face_12 = c(0.06,-0.11,0.95,0.7,0.22,0.89,-0.11,-0.89,-0.19,-0.22,0.5,0.56,0.29,-0.03,-0.29,0.96,-0.4,0.48,0.15,0.03,-0.78,-0.41,0.32,-0.84,-0.11,1.33,0.45,-0.28,0.38,0.57,-0.89,-1.04,-0.01,-0.42,-1.19,0.95,-0.43,-0.7,-0.32,-0.95),
	face_13 = c(0.22,-0.05,0.49,0.17,0.11,0.24,-0.21,-0.1,-0.05,-0.53,0.22,-0.35,0.07,0.5,-1.02,-0.32,0.81,-0.38,-0.4,0.91,-0.08,-0.05,-0.95,-0.45,-0.42,-0.35,-0.1,0.49,0.08,0.15,-0.49,0.44,0.07,-0.1,0.15,-0.14,0.41,0.35,-0.21,0.68),
	face_14 = c(-0.26,-0.73,1.47,-0.61,0,-1.33,0.61,0.56,0.22,-0.03,0.57,-0.16,-0.03,0.33,0.39,0.18,-0.43,0.4,-0.44,-0.63,0.02,0.21,-0.42,-1.11,-0.56,1.11,-0.56,-0.18,0.89,0.11,-0.72,-0.81,-1.59,0.37,0.64,0.28,-0.36,-0.93,-0.09,-1.08),
	face_15 = c(0.89,0.44,0.31,-0.39,-0.15,0.25,0.25,0.27,-0.82,0.27,0.14,-0.99,0.26,0.43,0.07,0.16,-0.14,-1.23,0.87,0.7,0.12,0.55,-0.82,0.85,-0.57,-0.76,0.21,0.77,-0.59,0.18,0.26,0.52,0.5,-0.24,-0.51,0.16,0.18,1.17,-0.16,0),
	face_16 = c(0.29,-0.71,0.72,0.79,-0.21,-0.14,-0.21,0.12,-0.29,1.03,0.04,-0.86,0.44,0.75,-0.39,-0.17,0.24,-0.7,0.36,0.27,0.31,-0.4,0.22,-0.52,0.28,-0.23,0.97,0.06,-0.38,-0.34,0.53,-0.64,0.64,-0.32,-0.01,0.67,-1.35,0.9,0.5,-0.84),
	face_17 = c(-0.17,0.03,-0.62,0.01,0.48,0.53,-0.6,-0.37,0.46,0.22,-0.36,-0.56,-0.25,0.47,0.59,-0.16,0.3,-0.29,0.04,1.43,0.12,-0.45,0.4,0.21,-0.09,-0.6,-0.24,0.8,0.38,0.14,-0.02,1.2,1.2,0.26,0.04,0.04,-0.13,-0.09,0.13,1.09),
	face_18 = c(0.06,-0.61,-0.86,0.62,0.37,-0.28,0.35,-0.66,0.64,-0.18,1.63,-0.13,-0.33,-0.62,0.24,0.13,0.04,-0.46,-0.8,0,0.33,-0.49,0.9,-0.38,-0.59,-0.24,-0.55,0.26,0.29,0.27,0.35,0.37,-0.28,0.76,0.28,-0.8,-0.24,0.81,0.29,-0.2),
	face_19 = c(0.3,0.3,-0.88,0.18,-1.17,0.56,-0.01,-0.01,-0.05,1.19,-0.03,0.01,0.24,-0.34,0.49,-0.16,0.22,0.26,0.37,1.21,0.45,0.3,0.59,0,-0.46,-0.68,-0.06,0.57,-0.77,0.82,0.94,0.9,0.55,0.02,-0.44,-0.02,0.17,0.83,-0.06,-1.12),
	face_20 = c(-0.08,-0.68,-0.37,-0.41,1.52,0.23,-0.94,-0.66,0.69,0.26,0.76,0.04,0.62,-0.42,0.02,0,-0.03,-0.43,0.55,0.06,0.9,-1.15,1.17,0.03,-0.11,-0.65,-0.25,-0.01,-1.06,0.13,0.18,0.76,1.85,1.49,0.2,0.84,0.46,-0.55,-0.59,0.07),
	face_21 = c(-0.76,-0.32,1.08,-0.25,0.46,-0.8,0.75,0.51,-0.21,-1.85,1.41,-0.52,-0.32,-0.26,-0.13,0.62,-0.61,0.13,0.17,-0.47,-0.91,0.41,0.57,-1.47,-0.12,0.35,-0.72,-0.19,0.9,-0.29,-2.02,-1.12,-1.71,0.56,1.11,0.45,-0.34,-1.33,-0.94,-0.98),
	face_22 = c(0.38,0.67,-0.13,0.36,1.22,0.45,-0.85,-0.28,-0.48,0.7,-0.8,0.49,0.05,-0.53,0.71,-0.34,0.63,0.26,-0.75,0.87,-0.11,0.53,-0.52,-0.11,-0.19,-0.28,-0.17,-0.65,-1.05,0.61,0.81,1.05,0.55,0.38,-0.08,-0.27,-0.52,0.88,0.8,0.98),
	face_23 = c(-0.61,-0.74,0.25,-0.4,0.08,0.05,1.12,0.8,-0.2,0.48,-0.21,0.85,-0.42,-0.08,-0.25,0.48,0.7,-0.59,-1.19,-0.45,-0.57,-0.15,1.11,0.13,-0.34,-0.34,0,-0.07,-0.46,0.06,0.24,-0.73,-0.78,-1.04,0.48,-0.06,0.23,0.48,-0.79,-0.06),
	face_24 = c(0.16,-0.77,0.2,-0.59,-1.21,-0.32,0.65,-0.31,0.45,-0.11,-0.3,0.2,-0.25,0.2,0.33,-0.94,0.5,0.3,-0.62,0.09,0.35,0.31,0.05,0.05,0.33,-0.18,0.35,-0.91,-0.52,-0.07,0.03,0.96,-0.43,-0.8,-0.37,-0.21,0.11,0.81,0.22,-0.18),
	face_25 = c(-0.18,0.73,-0.21,0.22,-0.39,0.43,0.68,-0.27,-0.73,0.3,-0.08,-0.06,0.38,-0.02,0.46,0.44,0.38,-0.43,-0.53,0.01,0.85,-0.7,0.11,-0.17,-0.04,-0.42,0.48,-0.12,-1.06,0.49,0.65,0.93,-0.37,0.29,0.22,0.69,-0.91,-1,-0.05,0.23),
	face_26 = c(0.51,0.26,-0.77,0.24,0.39,0.61,-0.78,-0.94,0.08,0.09,-1.08,-0.33,-1.51,-0.14,0.34,-0.32,0.25,0.19,-0.02,-0.38,-0.26,-0.33,-0.2,-0.22,-0.14,-0.08,1.36,0.28,-0.77,-0.98,-0.07,0.44,0.66,0.63,-0.59,-0.54,0.09,0.58,0.59,0.64),
	face_27 = c(-0.94,0.17,-0.81,0.93,0.01,0.72,0.04,0.33,0.01,-0.22,0.25,-0.83,0.03,0.88,-0.21,-0.56,-0.6,0.19,0.57,0.46,-0.43,-0.23,-0.04,-0.77,0.97,-0.92,-0.66,0.25,0.33,-0.35,1.01,0.39,0.09,0.05,-0.81,-0.12,0.65,0.54,-0.31,0.12),
	face_28 = c(0.04,0.43,-0.85,0.12,1.23,0.93,0.08,0.42,-0.41,1.22,0,-0.1,-0.3,0.48,0.48,0.31,0.52,-0.56,0.44,0.05,0.34,-0.04,-0.33,0.05,0.68,-0.65,0.95,0.03,-0.24,-0.09,-0.32,-0.32,0.76,0.15,-0.5,-0.16,0.56,-0.45,0.68,0.32),
	face_29 = c(0.71,-0.12,-0.14,-0.15,0.86,0.13,-0.54,-0.9,-0.52,0.04,-0.18,0.16,-0.01,0.47,-0.49,-0.01,-0.52,-0.38,0.46,1.11,0.17,0.21,0.11,0.35,-0.12,0.15,0.17,0.33,0.46,-0.79,0.31,1.11,-0.35,-0.44,0.66,-0.61,-0.69,1.08,-0.6,0.1),
	face_30 = c(0.14,-0.39,-1.6,0.59,0.3,0.81,0.1,-0.89,-0.03,0.26,-0.54,-0.14,0.26,-0.31,-0.14,0.89,-0.5,0.17,0.62,0.64,-0.21,0.34,-0.02,-0.64,0.25,-1.08,0.13,-0.23,-0.95,-0.75,0.64,0.66,1.63,0.18,-1.11,-0.65,-0.93,1.05,0.67,1.52))
```

## Introduction

This is the last of the weekly exercises, congrats on getting through them all! These exercises are designed to reinforce content taught each week. Typically, they consist of quiz questions and R exercises that require you to enter code into a box. These exercises are usually split into beginner and advance sections; however, this week all exercises are advanced. It is expected that students can complete the beginner sections of each weekly exercise, with the advanced sections designed for students who want an extra challenge or learn additional functions in R. If you are struggling, each exercise comes with a number of hints that may be helpful.

## Basic Exercises

This week's exercises cover linear mixed effect models, which is quite an advanced statistical technique. As such, all exercises this week can be considered "Advanced". You only need to complete them if you wish to practice conducting linear mixed effects models in R/love statistics so much.

## Advanced Exercises

### Question 1

The following code has a formula where the variable `var1` predicts the variable `DV`. Add a random intercept for a variable called `group1` to this formula.

```{r question_1-setup}

```

```{r question_1, exercise=TRUE,exercise.lines = 4}
formula <- 'DV ~ var1'
```

```{r question_1-hint-1}
Add random effects to a formula using the `+` symbol.
```

```{r question_1-hint-2}
Random effects are specified within brackets with the `|` symbol.

Next hint will show the solution...
```

```{r question_1-solution}
formula <- 'DV ~ var1 + (1 | group1)'
```

```{r question_1-code-check}
grade_code()
```

### Question 2

Answer the following question about the code below.

```{r question_2-setup}

```

```{r question_2, exercise=TRUE,exercise.lines = 4}
formula <- 'humour ~ wit + emotional_intelligence +
  (1 + wit + emotional_intelligence | country)'
```

```{r question_2-quiz}
quiz(question("Select all the correct statements about the formula above:",
     answer("The outcome variable (DV) is humour.",correct = TRUE),
     answer("There are two fixed effects.",correct = TRUE),
     answer("There is a random intercept of country.",correct = TRUE),
     answer("There are two random slopes for wit and emotional intelligence.",correct = TRUE),
     answer("There are two random intercepts for wit and emotional intelligence."),
     answer("Country is a fixed effect."),
     answer("No random intercept is specified."),
     answer("Emotional intelligence is more important than wit when predicting humour."),allow_retry = TRUE,random_answer_order = TRUE))
```

### Question 3

You are running a person perception study, where groups of participants interact and make judgements about each other. Below is a formula where participants judgements of trustworthiness of another participant is predicted by the rater's number of siblings. Currently, random intercepts are specified for participant and group, but no random slopes are specified. Edit the code the maximally include random slopes for number of siblings.

```{r question_3-setup}

```

```{r question_3, exercise=TRUE,exercise.lines = 4}
formula <- 'trustworthiness ~ siblings +
  (1 | participant_id) +
  (1 | group_id)'
```

```{r question_3-hint-1}
Add random slopes within the brackets on the left hand side of the `|` symbol.
```

```{r question_3-hint-2}
Number of siblings does not vary within participants, therefore, you do not have to specify a random slope of siblings in participant_id.

Next hint will show the solution...
```

```{r question_3-solution}
formula <- 'trustworthiness ~ siblings +
  (1 | participant_id) +
  (1 + siblings | group_id)'
```

```{r question_3-code-check}
grade_code()
```

### Question 4

You are testing children's language development and recruit 30 children from 30 schools (900 participants in total) across Scotland. You wish to test whether the child's family socioeconomic status influences language development, but control for the wealth of the area of the school in the model. Create a lmer formula to test this hypothesis. Specify a random intercept for school, and random slopes maximally. For the variable names, see the data.frame `language.data`.

```{r}
language.data <- data.frame(school_id = c(1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30),
							child_id = c(11,13,16,51,106,145,161,182,291,297,305,318,340,353,415,464,479,494,521,560,578,623,665,704,708,709,800,836,843,894,3,9,30,37,40,64,135,141,143,163,184,220,246,317,330,331,360,380,391,433,446,452,462,471,550,729,765,771,792,884,16,39,57,61,75,99,143,161,165,211,255,281,285,338,344,351,354,359,392,469,473,488,585,734,782,812,828,837,855,896,6,14,42,53,65,85,130,141,174,178,183,198,220,242,282,293,301,326,343,345,593,594,595,603,605,625,743,746,832,853,36,47,84,86,112,113,131,162,228,293,330,358,434,449,478,505,507,509,542,597,602,603,661,676,711,732,737,758,802,840,2,26,49,54,69,81,98,139,150,156,170,289,392,466,486,495,513,532,565,582,590,658,661,674,687,769,770,780,864,892,2,61,177,183,201,226,272,365,423,441,445,482,492,512,523,533,548,582,594,664,737,765,768,810,842,849,850,855,863,873,67,78,88,101,108,169,172,179,184,232,237,239,270,355,368,390,496,542,565,591,594,647,651,659,664,697,719,830,864,877,17,27,43,72,88,114,126,251,275,291,307,323,350,370,399,421,545,572,573,575,590,595,612,654,657,761,792,831,863,886,29,36,144,198,210,221,250,253,316,362,374,394,429,475,482,511,542,577,580,584,594,627,637,685,694,724,779,818,820,846,13,44,66,70,72,73,125,135,194,272,352,464,468,479,492,506,546,566,573,587,603,610,619,626,645,686,714,781,834,865,3,37,41,68,83,191,215,221,238,256,269,278,282,306,324,365,387,405,424,465,525,530,544,560,602,604,628,646,801,868,5,26,30,61,64,119,134,168,185,301,328,334,378,443,450,475,536,542,551,583,608,640,655,703,731,754,768,775,813,876,7,64,92,100,104,123,149,155,172,246,256,293,317,320,442,485,519,583,617,639,654,685,726,744,764,771,785,791,792,814,12,89,123,129,140,204,222,233,243,263,298,315,389,414,451,464,466,535,576,617,618,643,742,756,771,785,818,819,867,877,48,85,96,174,184,199,305,318,323,363,374,424,463,474,497,528,581,584,588,605,624,627,672,739,754,798,836,892,898,900,1,77,96,115,140,146,150,201,209,221,223,259,364,377,393,465,467,494,500,513,536,559,608,620,685,741,768,816,826,886,5,11,12,82,104,133,152,153,165,169,204,226,276,306,347,373,483,496,520,560,565,577,702,725,731,785,797,825,865,870,13,15,22,130,137,158,167,193,197,203,224,229,233,259,318,354,408,418,432,473,480,527,538,629,663,693,727,770,888,896,5,129,194,203,209,233,236,257,259,260,278,283,305,357,371,382,397,404,514,525,548,578,616,646,693,710,744,798,819,822,29,60,77,85,143,162,216,264,304,312,349,352,362,390,408,415,434,460,490,493,624,627,647,673,741,803,853,856,862,870,1,17,18,51,63,68,284,316,326,360,371,429,440,503,525,538,544,545,571,593,624,674,690,696,758,761,772,785,803,812,10,13,81,86,88,115,150,209,244,247,267,286,292,301,340,379,448,460,473,491,494,496,553,653,678,706,764,774,835,837,76,137,230,242,250,278,284,352,366,408,415,427,482,498,526,532,542,555,566,567,575,583,632,644,684,760,809,827,870,900,16,33,53,74,228,245,265,272,281,283,292,302,305,306,402,436,444,485,494,501,506,552,571,724,753,775,836,854,879,895,53,57,60,105,106,117,166,167,234,264,320,449,507,514,563,572,610,650,655,678,722,738,742,747,753,757,771,794,856,900,3,42,54,55,63,124,125,167,199,224,237,351,358,363,438,441,446,449,463,554,636,676,701,741,771,776,821,822,855,882,12,45,56,113,118,160,221,239,273,275,283,337,411,414,441,482,486,492,619,635,676,683,695,701,755,761,783,806,849,867,45,47,75,88,91,126,192,243,253,287,309,367,404,447,461,465,487,566,570,597,605,660,675,695,728,733,760,777,836,893,6,16,43,77,91,106,130,151,176,182,196,327,346,356,393,403,409,462,490,507,508,534,577,599,606,671,696,723,862,870),
							language = c(0.19,-1.07,0.43,-0.98,0.26,0.53,0.21,0.36,-0.53,0.21,-1.01,0.55,0.85,0.21,-0.76,-0.74,-0.59,-1.18,-0.36,0.4,-0.55,0.27,1.05,0.45,0.8,0.4,0.67,0.11,-1.03,-0.9,-0.89,0.52,-1.15,0.14,0.62,-0.08,0.45,0.43,-0.4,-0.26,0.51,-0.01,-0.09,0.59,0.6,-0.04,-0.97,0.45,-1.19,-0.53,0.53,0.1,0.52,0.05,0.28,-0.61,0.05,-0.25,-0.63,0.35,0.2,0.4,-0.41,-0.05,0.16,0.32,-0.58,-0.21,-0.29,-0.2,-1.61,-0.13,-0.4,-0.61,0.01,-0.7,0.15,-0.34,-0.84,-0.64,-0.29,0.36,-0.67,-0.76,-0.94,0.22,-0.1,0.61,0.73,-0.57,0.9,0.28,-0.63,-0.21,-0.52,-0.65,0.07,-0.19,-0.25,1.14,0.77,-0.26,-0.11,-0.42,0.77,-0.1,-0.11,0.24,-0.07,0.61,0,-1.39,-0.71,0.45,0.06,-0.11,0.36,0.01,0.72,-0.49,1.25,-0.22,-0.02,0.74,-0.53,1.05,-1.03,0.02,-0.1,-0.31,0.51,0,1.02,0.99,0.66,-0.52,0.06,0.3,-0.52,0.59,-0.81,0.5,-0.11,0.06,0.48,-1.05,0.43,1.5,0.32,-0.41,0.22,1.02,-0.2,-0.98,-0.08,0.48,-0.29,0.44,-0.13,-0.2,0.13,1.2,0.2,0.05,0.7,0.52,-0.15,-0.85,0.7,0.45,0.83,0.3,-1.02,0.08,0.62,-0.06,-0.14,0.08,0.03,-0.45,-0.65,-0.77,-0.67,0.38,-0.53,-0.45,0.06,0.03,-0.15,0.55,0.07,-1.91,-0.27,-1.16,-0.3,-0.5,-1.19,-0.8,-1.02,-0.86,-0.57,-0.41,-0.13,-0.58,-0.35,-0.15,-0.98,-0.72,0.01,-0.86,-0.5,0.56,-0.01,0.05,-0.84,0.87,0.65,0.03,0.11,0.77,-0.09,-0.21,0.06,0.3,0.47,-1.05,-0.7,-0.79,0.48,-0.42,-1.1,-0.17,0.86,-0.66,0.44,0.74,0.75,0.29,0.11,0.74,0.17,-0.42,-0.48,-0.12,-0.08,0.28,0.8,0.13,-0.13,-0.36,0.78,-0.13,0.17,-0.12,-1.87,0.84,0.93,0.1,-0.79,-0.23,0.41,0.04,-0.87,-0.42,-0.25,0.23,-0.41,-0.3,0.12,-0.5,-0.91,0.41,-0.09,-0.48,0.25,-0.03,0.6,-0.09,-0.53,0.1,0.16,0.47,-1.23,-1.15,-1.72,0.45,-0.27,-0.21,0.35,-0.02,-0.3,0.26,0.17,0.21,0.71,-0.18,-0.02,0.35,-0.32,-0.3,-0.95,-0.17,-0.45,-1.54,0.22,0.05,0.58,-0.31,-0.45,-0.35,-0.46,-0.78,-0.36,-0.87,0.49,1.31,-1,0.37,-0.34,-0.26,0.56,-0.74,-0.01,-0.97,0.24,0.49,-0.45,-0.39,-0.09,-1.01,-0.21,0.47,0.27,0.14,0.01,-0.08,0.23,-0.5,0.05,-0.88,0.18,0.22,0.31,-0.79,0.18,1.08,1.14,0.05,0.61,-0.22,-0.35,0.6,-0.82,-0.34,0.22,-1.34,-0.14,0.56,0.18,-0.51,1,-0.45,-0.48,0.19,-0.46,-0.08,0.8,-0.86,0.39,-0.24,0.09,1.17,0.51,-0.58,0.57,-0.33,-0.2,0.5,0.82,0.49,-0.49,-0.43,-0.03,0.2,-0.58,1.82,1.16,0.82,-0.02,0.18,0.76,-0.13,-1.48,0.51,0,0.16,1,-0.27,0.72,-0.31,0.51,0.45,-0.35,-0.62,-0.34,0.44,0.35,-0.21,0.62,-1.24,-0.01,-0.3,-0.28,-0.49,0.85,-0.98,0.82,-0.17,-0.11,0.84,-0.42,-0.74,-0.41,-0.14,0.68,-0.79,-0.65,0.03,-1.2,-0.19,0.7,0.54,0.23,-0.7,-0.44,-0.56,0.16,0.01,-1.19,-1.38,0.8,-0.19,0.58,-0.66,-0.6,0.42,-0.52,-0.14,0.37,0.55,0.13,-0.04,0.09,0.13,0.16,0.74,-0.23,0.64,1.04,-0.01,0.65,0.69,0.02,0.24,0.56,0.85,0.09,-0.46,1.02,0.94,0.38,-0.4,0.32,0.01,0.64,1.35,-0.07,-0.71,-0.8,0.08,0.11,-0.9,-0.62,0.37,-0.18,0.18,-0.14,-0.74,0.96,0.5,-0.13,-0.7,-0.21,1.9,0.37,0.46,-0.13,-0.3,-0.22,-0.15,0.16,-0.35,0.18,-0.17,-0.57,-0.44,0.74,0.17,0.3,0.04,-0.22,0.32,0.68,0.29,-0.1,-1.11,-0.97,-0.34,-0.7,1.14,-0.89,0.32,-0.75,0.07,-0.51,-0.11,0.02,-1.38,-1.53,-0.72,-0.42,0.31,-0.42,0.21,-0.02,0.38,1.02,0.18,-0.56,-0.83,-1.56,-0.77,0.08,-0.14,-0.81,-0.75,-1.03,0.11,-1.37,-0.84,-0.6,-0.74,-0.48,0.3,0.87,0.22,-0.69,0.33,-0.42,-0.26,-0.53,0.11,-0.95,-1.47,-0.9,-0.6,-0.5,-1.34,1,0.28,-0.33,-0.02,-0.17,-0.74,0.53,-0.32,0.37,0.16,-0.8,0.07,-0.66,0.1,-0.42,-0.7,0.01,-0.26,-0.31,0.5,0.81,0.44,-0.49,-0.11,-0.58,-0.52,-0.92,-0.12,0.44,0.21,0.06,-0.44,-0.72,-0.92,-0.77,-0.71,-0.32,-0.07,-0.43,-0.45,-0.38,0.22,0.11,-0.58,0.44,0.07,-0.49,-0.88,0.54,0.65,-0.46,-0.03,0.48,0.08,-0.11,0.97,0.23,-1.13,-1.2,-0.17,-1.31,-0.52,0.09,0.18,0.09,-0.35,-0.8,1.01,-0.9,-0.66,0.12,-0.76,-0.87,0.99,0.39,-0.34,0.49,-0.19,-0.29,-0.47,0.86,-0.9,-1.18,0.91,0.01,-0.54,1.02,-0.01,-0.22,0.17,-1.3,-0.42,-0.19,-0.46,0.19,1,0.22,0.25,0.12,1.16,0.59,0.11,0.5,1.05,-0.29,0.39,0.38,-0.54,0.71,0.43,-0.07,0.03,0.31,-1.42,0.34,-0.43,-0.09,0.61,0.82,-0.64,0.45,0.08,0.4,0.27,-0.23,-0.23,0.66,-0.43,-1,-0.48,-0.39,0.75,0.54,-0.63,-0.18,-0.85,-0.69,0.35,0.49,0,-0.14,-0.05,-0.16,-0.53,0.24,-0.18,-0.88,-1.09,0.24,-0.42,0.57,-0.52,-0.26,0.38,-0.98,0.19,-0.19,-0.16,0.07,-1.53,-0.11,-0.15,0.42,-0.93,-0.24,0.22,-0.2,0.63,-0.31,0.41,-0.8,0.41,0.57,0.12,-0.32,0.27,0.97,0.58,1.21,0.55,0.65,-0.15,0.98,-0.71,-0.02,-0.53,0.24,0.97,0.14,-0.54,0.78,-0.29,-1.08,0.29,-0.56,0,1.03,-0.12,0.39,0.48,-0.21,-0.9,0.29,0.25,-0.39,0.54,-0.14,-0.21,-1.02,-0.21,0.38,-0.92,-2.46,-1.99,-1.17,-1.15,-1.27,0.72,-1.06,-1,-0.24,0.27,-0.4,-0.81,-0.54,-1.15,-0.88,0.11,-0.36,0.69,-0.73,-1.09,-1.53,-0.72,0.34,-0.05,-0.19,-0.74,-1.29,-0.88,-0.91,-0.64,-0.21,0.38,-0.35,0.06,-0.53,-0.28,-0.26,-0.13,-0.62,-1.25,0.43,0.77,0.1,-0.11,1.41,-0.69,0.34,-1.12,0.3,0.13,0.27,-0.11,0.79,0.64,-0.19,-0.66,-1.15,-0.78,0.78,-0.49,0.07,0.09,0.9,-1.06,0.47,0.72,0.21,0.28,0.48,-0.01,-0.52,0.13,0.46,0.3,0.2,0.47,0.4,-0.34,0.98,-0.06,0.87,-0.11,0.13,0.12,-0.24,-0.3,0.81,0.2,1.35,-0.27,1.09,0.26,0.38,-0.75,0.99,1.24,0.46,-0.33,-0.56,0.23,-0.22,0.1,0.13,0.14,1.38,0.01,0.24,-0.89,0.18,-0.06,-0.03,0.37,0.19,1.07,0.42,0.09,-0.23,0.38,-0.36,-0.08),
							family_ses = c(0.44,-2.16,0.29,-1.17,-0.8,-0.25,0.64,-1.37,-0.35,0.95,0.21,2.56,2.4,0.47,-0.56,-1.54,-1.06,-1.89,0.86,-1.16,0.59,-0.04,1.27,0.95,0.53,0.39,0.36,2.91,0.82,-0.45,-3.16,1.56,-0.54,-0.15,0.12,-0.01,0.91,0.98,-1.02,-1.34,0.92,-0.65,-0.36,-0.14,1.72,0.05,-1.56,-0.1,-0.95,-1.66,0.63,-0.73,1.07,-0.11,1.06,0.96,0.46,-1.12,-1.22,1.16,0.29,-0.93,0.46,-0.03,1.87,0.69,-1.02,0.64,-2.45,0.74,-0.25,1.42,0.29,1.14,-0.22,-0.58,-0.94,-1.07,-1.58,-2.05,0.24,1.25,-1.3,-0.34,-1.82,0.18,0.77,-0.31,-0.35,-1.01,0.88,-0.05,-4.24,0.95,-1.53,-1.71,-0.68,0.98,-0.11,0.79,1.98,0.51,-0.65,-0.52,-0.42,-0.66,-0.31,0.54,1.07,0.64,0.51,-1.68,-1.14,0.36,1.59,0.73,0.81,0.88,0.48,-1.56,1.18,-1.16,0.15,1.25,-0.87,-0.56,-1.45,-0.2,-0.09,-0.66,1.72,-0.28,0.71,2.25,0.85,-1.34,-0.02,0.47,0.22,1.38,-0.81,0.36,0.19,-0.44,0.81,-0.73,0.26,1.66,0.89,-1.5,0.17,-0.42,-0.91,-1.83,0.17,-0.61,1.35,0.48,0.82,-0.41,0.2,0.7,-1.58,0.73,0.87,0.98,-0.75,-1.06,0.38,0.44,1.65,0.39,0.19,1.26,0.09,-0.77,-1.32,-2.62,0.17,-1.01,0.17,-0.03,-0.73,1.98,0.08,-0.36,-0.62,2.07,0.69,2.02,-0.02,-1.48,-0.54,-1.42,-0.49,-0.02,-1.24,0.44,-1.68,0.52,0.26,0.46,-0.45,1.75,0.38,1.6,-0.66,-0.35,0.89,-0.06,0.26,0.3,0.25,1.25,-1.49,1.88,1.14,1.57,0.92,1.11,-0.32,-0.05,0.72,2.17,0.75,-1.7,-1.23,0.22,0.38,0.6,-1.68,-0.11,-0.8,-0.63,0.52,1.48,0.47,0.05,0.17,0.54,0.09,-0.36,0.83,-0.36,0.25,0.37,0.03,0.34,-1.18,-0.35,-0.71,1.33,0.36,-1.14,-1.93,1.71,-0.19,0.17,-0.39,0.69,1.65,-1.14,0.67,1.06,-1.87,-0.9,-1.22,-1.31,0.89,-0.15,-0.51,1.18,-0.38,0.51,-0.42,-1.14,0.59,0.69,-0.99,0.57,-0.21,2.65,-0.86,-1.46,-1.48,0.38,0.22,0.28,0.92,-0.22,-1.68,2.08,-0.88,-1.65,0.08,0.37,0.76,-0.57,-0.49,-0.78,-2.16,0.06,-0.38,-0.68,-0.36,1.14,-0.08,0.91,0.1,-0.62,-0.62,-1.54,0.48,-1.06,-0.54,1.56,0.05,0.22,-0.39,-0.53,0.36,0.11,0.09,-0.17,0.74,0.81,1.75,1.26,1.44,-0.83,-3.16,-0.15,1.1,-0.03,1.07,-1.6,0.04,-1.14,-0.12,-0.63,-0.28,-1.25,-0.42,0.13,-0.53,2.07,0.78,0.52,3.19,0.49,1.09,0.87,0.02,-1.16,-0.81,-0.95,-0.18,-0.97,0.41,-0.4,-0.57,-0.42,-0.54,-0.03,-0.01,-1.08,1.39,0.54,0.65,-0.31,-0.41,0.17,-1.35,-0.05,0.57,-1.46,-0.97,0.22,0.59,-0.52,-0.13,-1.22,-0.26,-0.17,0.03,1.21,-0.45,0.84,-1.04,0.97,1.08,-0.01,-1.67,0.17,-0.14,0.75,1.75,0.08,1.14,-0.36,-0.63,-0.66,-0.14,-1.23,0.48,1.44,0.07,-0.52,-0.76,-0.82,1.06,-1.65,-1.05,0.72,1.56,-1.12,1.7,0.25,-1.22,2.18,0.32,-1.45,0.75,0.76,-0.28,-0.01,-0.61,0.26,-0.85,-1.75,0.24,-1.79,-1.79,-0.65,-0.21,-1.54,0.73,0.83,0.83,-0.76,1.63,1.32,0.05,-1.66,-1.12,1.7,-0.57,1.22,0.4,0.54,-0.55,-1.71,-1,-0.11,0.92,-0.39,0.21,2.56,1.33,-1.82,-0.21,3.19,-1.07,-0.54,1.49,-0.08,-1.89,-0.22,0.92,1.59,0.05,2.08,1.3,0.69,1.21,0.11,2.91,-1.01,-2.62,-0.05,0.18,-1.69,-1,0.71,-0.28,-0.87,0.82,0.08,2.48,-1.14,0.56,-0.34,0.23,0.51,1.64,0.49,-0.15,-1.89,1.22,-0.75,-0.97,0.48,-0.13,-0.05,-1.65,0.27,-0.45,1.46,0.28,-0.15,-0.57,0.44,0.32,-1.06,-0.14,-0.8,-0.3,-0.09,-2.45,1.88,-0.01,-0.36,0.11,0.13,-0.93,-0.82,1.54,-1.23,-1.72,-1.16,0.38,0.28,-2.11,0.51,0.03,1.7,1.71,0.87,-0.83,-0.21,-2.16,-0.47,0.95,-0.68,0.31,0.41,-1.58,0.6,-2.11,-1.48,0.86,0.58,0.26,-0.34,2.56,-0.94,-0.48,0.65,-0.91,0.24,1.06,1.49,-0.12,0.17,-2.43,-2.39,-0.56,-1.32,0.41,-1.01,-0.57,0.76,0.1,-1.48,2.48,0.26,-0.36,1.07,-0.34,-1.03,-1.25,1.24,0.21,-0.25,-0.72,-0.18,1.11,0.59,1.31,1.09,-1.24,0.59,0.83,-0.97,-2.39,0.37,0.72,0.11,1.22,-1.28,-0.51,-0.08,-1.69,-1.71,-1.02,-0.2,-0.18,0.2,-0.96,0.77,-0.14,-0.62,0.57,-1.7,-0.48,-0.56,0.71,0.19,0.1,-0.16,0.05,2.08,-0.11,1.74,0.27,-0.89,-1.56,0.79,-0.73,-0.21,0.18,0.09,-1.04,-1.17,-1.06,-0.03,-1.75,-0.99,0.54,-1.56,-0.72,-0.86,2.09,-0.01,1.09,-0.12,0.02,-0.19,0.63,0.51,0.05,1.26,-0.77,-0.27,1.66,-0.9,-0.66,1.7,-0.89,0.18,1.32,-2.16,-0.61,1.25,0.25,0.71,0.82,2.48,1.44,-1.95,-0.16,0.23,0.46,-0.31,2.4,0.11,1.34,0.19,0.24,0.49,-1.89,-1.23,-0.84,-0.86,-1.46,0.6,1.56,-1.27,0.21,-0.31,-0.18,0.31,0.01,-0.52,0.59,-1.25,-1.75,-0.62,0.04,-0.48,-0.56,1.31,-1.48,-0.06,0.74,-1.06,0.22,-0.69,0.22,-0.19,0.69,-0.52,0.63,-0.58,-1.32,-1.28,-0.27,-0.99,-0.21,-0.05,0.29,1.33,0.95,0,-0.09,-0.28,0.17,-0.62,1.42,1.24,0.46,-0.9,0.21,0.13,-0.2,0.24,0.08,1.44,-1.89,1.57,1.56,1.12,0.63,0.37,0.19,0.84,2.91,1.34,1.17,-1.12,0.95,0.46,-0.08,-0.64,-0.8,1.83,0.2,-1.58,0.58,0.2,-1.23,2.25,-0.02,1.31,1.97,0.17,0.11,0.84,-0.26,-1.46,0.37,-0.3,0.05,0.31,0.19,0.86,-1.12,-0.62,0.79,-0.05,-3.16,-4.24,-1.83,-0.62,-1.06,1.11,-0.08,-1.58,-0.39,0.86,-0.32,-0.58,-0.28,-1.82,-2.27,2.02,0.63,2.25,-1.07,-0.2,-1.3,-0.44,1.84,0.27,-1.12,-0.39,-0.47,-1.28,-0.35,-0.71,0.32,1.05,-1.19,-0.56,-0.88,-1.76,-1.14,-0.05,-0.13,-1.18,1.24,1.43,0.59,-0.65,2.02,-1.48,0.87,-0.54,0.09,0.91,-0.44,-0.53,1.35,1.84,-0.32,-0.9,-0.3,-0.87,1.6,0.4,1.05,-1.16,1.87,0.25,0.51,0.03,0.53,-0.85,0.69,0.92,-1.53,-0.89,0.59,-0.53,-1.68,0.49,0.98,0.22,1.61,1.38,1.59,-0.33,0.79,1.35,-1.19,-1.27,-1.28,1.88,2.91,-0.29,0.88,0.29,0.83,-1.69,0.51,-0.8,-0.68,-0.75,-1.45,-1.37,-0.06,0.55,-0.77,-0.29,1.64,-0.79,-0.63,1.07,0.1,-0.02,-0.52,-1.21,0.28,1.56,0.4,0.43,-0.27,0.36,-0.73,-0.21),
							area_wealth = c(-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1))
							
language.data
```

```{r question_4-setup}
language.data <- data.frame(school_id = c(1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30),
							child_id = c(11,13,16,51,106,145,161,182,291,297,305,318,340,353,415,464,479,494,521,560,578,623,665,704,708,709,800,836,843,894,3,9,30,37,40,64,135,141,143,163,184,220,246,317,330,331,360,380,391,433,446,452,462,471,550,729,765,771,792,884,16,39,57,61,75,99,143,161,165,211,255,281,285,338,344,351,354,359,392,469,473,488,585,734,782,812,828,837,855,896,6,14,42,53,65,85,130,141,174,178,183,198,220,242,282,293,301,326,343,345,593,594,595,603,605,625,743,746,832,853,36,47,84,86,112,113,131,162,228,293,330,358,434,449,478,505,507,509,542,597,602,603,661,676,711,732,737,758,802,840,2,26,49,54,69,81,98,139,150,156,170,289,392,466,486,495,513,532,565,582,590,658,661,674,687,769,770,780,864,892,2,61,177,183,201,226,272,365,423,441,445,482,492,512,523,533,548,582,594,664,737,765,768,810,842,849,850,855,863,873,67,78,88,101,108,169,172,179,184,232,237,239,270,355,368,390,496,542,565,591,594,647,651,659,664,697,719,830,864,877,17,27,43,72,88,114,126,251,275,291,307,323,350,370,399,421,545,572,573,575,590,595,612,654,657,761,792,831,863,886,29,36,144,198,210,221,250,253,316,362,374,394,429,475,482,511,542,577,580,584,594,627,637,685,694,724,779,818,820,846,13,44,66,70,72,73,125,135,194,272,352,464,468,479,492,506,546,566,573,587,603,610,619,626,645,686,714,781,834,865,3,37,41,68,83,191,215,221,238,256,269,278,282,306,324,365,387,405,424,465,525,530,544,560,602,604,628,646,801,868,5,26,30,61,64,119,134,168,185,301,328,334,378,443,450,475,536,542,551,583,608,640,655,703,731,754,768,775,813,876,7,64,92,100,104,123,149,155,172,246,256,293,317,320,442,485,519,583,617,639,654,685,726,744,764,771,785,791,792,814,12,89,123,129,140,204,222,233,243,263,298,315,389,414,451,464,466,535,576,617,618,643,742,756,771,785,818,819,867,877,48,85,96,174,184,199,305,318,323,363,374,424,463,474,497,528,581,584,588,605,624,627,672,739,754,798,836,892,898,900,1,77,96,115,140,146,150,201,209,221,223,259,364,377,393,465,467,494,500,513,536,559,608,620,685,741,768,816,826,886,5,11,12,82,104,133,152,153,165,169,204,226,276,306,347,373,483,496,520,560,565,577,702,725,731,785,797,825,865,870,13,15,22,130,137,158,167,193,197,203,224,229,233,259,318,354,408,418,432,473,480,527,538,629,663,693,727,770,888,896,5,129,194,203,209,233,236,257,259,260,278,283,305,357,371,382,397,404,514,525,548,578,616,646,693,710,744,798,819,822,29,60,77,85,143,162,216,264,304,312,349,352,362,390,408,415,434,460,490,493,624,627,647,673,741,803,853,856,862,870,1,17,18,51,63,68,284,316,326,360,371,429,440,503,525,538,544,545,571,593,624,674,690,696,758,761,772,785,803,812,10,13,81,86,88,115,150,209,244,247,267,286,292,301,340,379,448,460,473,491,494,496,553,653,678,706,764,774,835,837,76,137,230,242,250,278,284,352,366,408,415,427,482,498,526,532,542,555,566,567,575,583,632,644,684,760,809,827,870,900,16,33,53,74,228,245,265,272,281,283,292,302,305,306,402,436,444,485,494,501,506,552,571,724,753,775,836,854,879,895,53,57,60,105,106,117,166,167,234,264,320,449,507,514,563,572,610,650,655,678,722,738,742,747,753,757,771,794,856,900,3,42,54,55,63,124,125,167,199,224,237,351,358,363,438,441,446,449,463,554,636,676,701,741,771,776,821,822,855,882,12,45,56,113,118,160,221,239,273,275,283,337,411,414,441,482,486,492,619,635,676,683,695,701,755,761,783,806,849,867,45,47,75,88,91,126,192,243,253,287,309,367,404,447,461,465,487,566,570,597,605,660,675,695,728,733,760,777,836,893,6,16,43,77,91,106,130,151,176,182,196,327,346,356,393,403,409,462,490,507,508,534,577,599,606,671,696,723,862,870),
							language = c(0.19,-1.07,0.43,-0.98,0.26,0.53,0.21,0.36,-0.53,0.21,-1.01,0.55,0.85,0.21,-0.76,-0.74,-0.59,-1.18,-0.36,0.4,-0.55,0.27,1.05,0.45,0.8,0.4,0.67,0.11,-1.03,-0.9,-0.89,0.52,-1.15,0.14,0.62,-0.08,0.45,0.43,-0.4,-0.26,0.51,-0.01,-0.09,0.59,0.6,-0.04,-0.97,0.45,-1.19,-0.53,0.53,0.1,0.52,0.05,0.28,-0.61,0.05,-0.25,-0.63,0.35,0.2,0.4,-0.41,-0.05,0.16,0.32,-0.58,-0.21,-0.29,-0.2,-1.61,-0.13,-0.4,-0.61,0.01,-0.7,0.15,-0.34,-0.84,-0.64,-0.29,0.36,-0.67,-0.76,-0.94,0.22,-0.1,0.61,0.73,-0.57,0.9,0.28,-0.63,-0.21,-0.52,-0.65,0.07,-0.19,-0.25,1.14,0.77,-0.26,-0.11,-0.42,0.77,-0.1,-0.11,0.24,-0.07,0.61,0,-1.39,-0.71,0.45,0.06,-0.11,0.36,0.01,0.72,-0.49,1.25,-0.22,-0.02,0.74,-0.53,1.05,-1.03,0.02,-0.1,-0.31,0.51,0,1.02,0.99,0.66,-0.52,0.06,0.3,-0.52,0.59,-0.81,0.5,-0.11,0.06,0.48,-1.05,0.43,1.5,0.32,-0.41,0.22,1.02,-0.2,-0.98,-0.08,0.48,-0.29,0.44,-0.13,-0.2,0.13,1.2,0.2,0.05,0.7,0.52,-0.15,-0.85,0.7,0.45,0.83,0.3,-1.02,0.08,0.62,-0.06,-0.14,0.08,0.03,-0.45,-0.65,-0.77,-0.67,0.38,-0.53,-0.45,0.06,0.03,-0.15,0.55,0.07,-1.91,-0.27,-1.16,-0.3,-0.5,-1.19,-0.8,-1.02,-0.86,-0.57,-0.41,-0.13,-0.58,-0.35,-0.15,-0.98,-0.72,0.01,-0.86,-0.5,0.56,-0.01,0.05,-0.84,0.87,0.65,0.03,0.11,0.77,-0.09,-0.21,0.06,0.3,0.47,-1.05,-0.7,-0.79,0.48,-0.42,-1.1,-0.17,0.86,-0.66,0.44,0.74,0.75,0.29,0.11,0.74,0.17,-0.42,-0.48,-0.12,-0.08,0.28,0.8,0.13,-0.13,-0.36,0.78,-0.13,0.17,-0.12,-1.87,0.84,0.93,0.1,-0.79,-0.23,0.41,0.04,-0.87,-0.42,-0.25,0.23,-0.41,-0.3,0.12,-0.5,-0.91,0.41,-0.09,-0.48,0.25,-0.03,0.6,-0.09,-0.53,0.1,0.16,0.47,-1.23,-1.15,-1.72,0.45,-0.27,-0.21,0.35,-0.02,-0.3,0.26,0.17,0.21,0.71,-0.18,-0.02,0.35,-0.32,-0.3,-0.95,-0.17,-0.45,-1.54,0.22,0.05,0.58,-0.31,-0.45,-0.35,-0.46,-0.78,-0.36,-0.87,0.49,1.31,-1,0.37,-0.34,-0.26,0.56,-0.74,-0.01,-0.97,0.24,0.49,-0.45,-0.39,-0.09,-1.01,-0.21,0.47,0.27,0.14,0.01,-0.08,0.23,-0.5,0.05,-0.88,0.18,0.22,0.31,-0.79,0.18,1.08,1.14,0.05,0.61,-0.22,-0.35,0.6,-0.82,-0.34,0.22,-1.34,-0.14,0.56,0.18,-0.51,1,-0.45,-0.48,0.19,-0.46,-0.08,0.8,-0.86,0.39,-0.24,0.09,1.17,0.51,-0.58,0.57,-0.33,-0.2,0.5,0.82,0.49,-0.49,-0.43,-0.03,0.2,-0.58,1.82,1.16,0.82,-0.02,0.18,0.76,-0.13,-1.48,0.51,0,0.16,1,-0.27,0.72,-0.31,0.51,0.45,-0.35,-0.62,-0.34,0.44,0.35,-0.21,0.62,-1.24,-0.01,-0.3,-0.28,-0.49,0.85,-0.98,0.82,-0.17,-0.11,0.84,-0.42,-0.74,-0.41,-0.14,0.68,-0.79,-0.65,0.03,-1.2,-0.19,0.7,0.54,0.23,-0.7,-0.44,-0.56,0.16,0.01,-1.19,-1.38,0.8,-0.19,0.58,-0.66,-0.6,0.42,-0.52,-0.14,0.37,0.55,0.13,-0.04,0.09,0.13,0.16,0.74,-0.23,0.64,1.04,-0.01,0.65,0.69,0.02,0.24,0.56,0.85,0.09,-0.46,1.02,0.94,0.38,-0.4,0.32,0.01,0.64,1.35,-0.07,-0.71,-0.8,0.08,0.11,-0.9,-0.62,0.37,-0.18,0.18,-0.14,-0.74,0.96,0.5,-0.13,-0.7,-0.21,1.9,0.37,0.46,-0.13,-0.3,-0.22,-0.15,0.16,-0.35,0.18,-0.17,-0.57,-0.44,0.74,0.17,0.3,0.04,-0.22,0.32,0.68,0.29,-0.1,-1.11,-0.97,-0.34,-0.7,1.14,-0.89,0.32,-0.75,0.07,-0.51,-0.11,0.02,-1.38,-1.53,-0.72,-0.42,0.31,-0.42,0.21,-0.02,0.38,1.02,0.18,-0.56,-0.83,-1.56,-0.77,0.08,-0.14,-0.81,-0.75,-1.03,0.11,-1.37,-0.84,-0.6,-0.74,-0.48,0.3,0.87,0.22,-0.69,0.33,-0.42,-0.26,-0.53,0.11,-0.95,-1.47,-0.9,-0.6,-0.5,-1.34,1,0.28,-0.33,-0.02,-0.17,-0.74,0.53,-0.32,0.37,0.16,-0.8,0.07,-0.66,0.1,-0.42,-0.7,0.01,-0.26,-0.31,0.5,0.81,0.44,-0.49,-0.11,-0.58,-0.52,-0.92,-0.12,0.44,0.21,0.06,-0.44,-0.72,-0.92,-0.77,-0.71,-0.32,-0.07,-0.43,-0.45,-0.38,0.22,0.11,-0.58,0.44,0.07,-0.49,-0.88,0.54,0.65,-0.46,-0.03,0.48,0.08,-0.11,0.97,0.23,-1.13,-1.2,-0.17,-1.31,-0.52,0.09,0.18,0.09,-0.35,-0.8,1.01,-0.9,-0.66,0.12,-0.76,-0.87,0.99,0.39,-0.34,0.49,-0.19,-0.29,-0.47,0.86,-0.9,-1.18,0.91,0.01,-0.54,1.02,-0.01,-0.22,0.17,-1.3,-0.42,-0.19,-0.46,0.19,1,0.22,0.25,0.12,1.16,0.59,0.11,0.5,1.05,-0.29,0.39,0.38,-0.54,0.71,0.43,-0.07,0.03,0.31,-1.42,0.34,-0.43,-0.09,0.61,0.82,-0.64,0.45,0.08,0.4,0.27,-0.23,-0.23,0.66,-0.43,-1,-0.48,-0.39,0.75,0.54,-0.63,-0.18,-0.85,-0.69,0.35,0.49,0,-0.14,-0.05,-0.16,-0.53,0.24,-0.18,-0.88,-1.09,0.24,-0.42,0.57,-0.52,-0.26,0.38,-0.98,0.19,-0.19,-0.16,0.07,-1.53,-0.11,-0.15,0.42,-0.93,-0.24,0.22,-0.2,0.63,-0.31,0.41,-0.8,0.41,0.57,0.12,-0.32,0.27,0.97,0.58,1.21,0.55,0.65,-0.15,0.98,-0.71,-0.02,-0.53,0.24,0.97,0.14,-0.54,0.78,-0.29,-1.08,0.29,-0.56,0,1.03,-0.12,0.39,0.48,-0.21,-0.9,0.29,0.25,-0.39,0.54,-0.14,-0.21,-1.02,-0.21,0.38,-0.92,-2.46,-1.99,-1.17,-1.15,-1.27,0.72,-1.06,-1,-0.24,0.27,-0.4,-0.81,-0.54,-1.15,-0.88,0.11,-0.36,0.69,-0.73,-1.09,-1.53,-0.72,0.34,-0.05,-0.19,-0.74,-1.29,-0.88,-0.91,-0.64,-0.21,0.38,-0.35,0.06,-0.53,-0.28,-0.26,-0.13,-0.62,-1.25,0.43,0.77,0.1,-0.11,1.41,-0.69,0.34,-1.12,0.3,0.13,0.27,-0.11,0.79,0.64,-0.19,-0.66,-1.15,-0.78,0.78,-0.49,0.07,0.09,0.9,-1.06,0.47,0.72,0.21,0.28,0.48,-0.01,-0.52,0.13,0.46,0.3,0.2,0.47,0.4,-0.34,0.98,-0.06,0.87,-0.11,0.13,0.12,-0.24,-0.3,0.81,0.2,1.35,-0.27,1.09,0.26,0.38,-0.75,0.99,1.24,0.46,-0.33,-0.56,0.23,-0.22,0.1,0.13,0.14,1.38,0.01,0.24,-0.89,0.18,-0.06,-0.03,0.37,0.19,1.07,0.42,0.09,-0.23,0.38,-0.36,-0.08),
							family_ses = c(0.44,-2.16,0.29,-1.17,-0.8,-0.25,0.64,-1.37,-0.35,0.95,0.21,2.56,2.4,0.47,-0.56,-1.54,-1.06,-1.89,0.86,-1.16,0.59,-0.04,1.27,0.95,0.53,0.39,0.36,2.91,0.82,-0.45,-3.16,1.56,-0.54,-0.15,0.12,-0.01,0.91,0.98,-1.02,-1.34,0.92,-0.65,-0.36,-0.14,1.72,0.05,-1.56,-0.1,-0.95,-1.66,0.63,-0.73,1.07,-0.11,1.06,0.96,0.46,-1.12,-1.22,1.16,0.29,-0.93,0.46,-0.03,1.87,0.69,-1.02,0.64,-2.45,0.74,-0.25,1.42,0.29,1.14,-0.22,-0.58,-0.94,-1.07,-1.58,-2.05,0.24,1.25,-1.3,-0.34,-1.82,0.18,0.77,-0.31,-0.35,-1.01,0.88,-0.05,-4.24,0.95,-1.53,-1.71,-0.68,0.98,-0.11,0.79,1.98,0.51,-0.65,-0.52,-0.42,-0.66,-0.31,0.54,1.07,0.64,0.51,-1.68,-1.14,0.36,1.59,0.73,0.81,0.88,0.48,-1.56,1.18,-1.16,0.15,1.25,-0.87,-0.56,-1.45,-0.2,-0.09,-0.66,1.72,-0.28,0.71,2.25,0.85,-1.34,-0.02,0.47,0.22,1.38,-0.81,0.36,0.19,-0.44,0.81,-0.73,0.26,1.66,0.89,-1.5,0.17,-0.42,-0.91,-1.83,0.17,-0.61,1.35,0.48,0.82,-0.41,0.2,0.7,-1.58,0.73,0.87,0.98,-0.75,-1.06,0.38,0.44,1.65,0.39,0.19,1.26,0.09,-0.77,-1.32,-2.62,0.17,-1.01,0.17,-0.03,-0.73,1.98,0.08,-0.36,-0.62,2.07,0.69,2.02,-0.02,-1.48,-0.54,-1.42,-0.49,-0.02,-1.24,0.44,-1.68,0.52,0.26,0.46,-0.45,1.75,0.38,1.6,-0.66,-0.35,0.89,-0.06,0.26,0.3,0.25,1.25,-1.49,1.88,1.14,1.57,0.92,1.11,-0.32,-0.05,0.72,2.17,0.75,-1.7,-1.23,0.22,0.38,0.6,-1.68,-0.11,-0.8,-0.63,0.52,1.48,0.47,0.05,0.17,0.54,0.09,-0.36,0.83,-0.36,0.25,0.37,0.03,0.34,-1.18,-0.35,-0.71,1.33,0.36,-1.14,-1.93,1.71,-0.19,0.17,-0.39,0.69,1.65,-1.14,0.67,1.06,-1.87,-0.9,-1.22,-1.31,0.89,-0.15,-0.51,1.18,-0.38,0.51,-0.42,-1.14,0.59,0.69,-0.99,0.57,-0.21,2.65,-0.86,-1.46,-1.48,0.38,0.22,0.28,0.92,-0.22,-1.68,2.08,-0.88,-1.65,0.08,0.37,0.76,-0.57,-0.49,-0.78,-2.16,0.06,-0.38,-0.68,-0.36,1.14,-0.08,0.91,0.1,-0.62,-0.62,-1.54,0.48,-1.06,-0.54,1.56,0.05,0.22,-0.39,-0.53,0.36,0.11,0.09,-0.17,0.74,0.81,1.75,1.26,1.44,-0.83,-3.16,-0.15,1.1,-0.03,1.07,-1.6,0.04,-1.14,-0.12,-0.63,-0.28,-1.25,-0.42,0.13,-0.53,2.07,0.78,0.52,3.19,0.49,1.09,0.87,0.02,-1.16,-0.81,-0.95,-0.18,-0.97,0.41,-0.4,-0.57,-0.42,-0.54,-0.03,-0.01,-1.08,1.39,0.54,0.65,-0.31,-0.41,0.17,-1.35,-0.05,0.57,-1.46,-0.97,0.22,0.59,-0.52,-0.13,-1.22,-0.26,-0.17,0.03,1.21,-0.45,0.84,-1.04,0.97,1.08,-0.01,-1.67,0.17,-0.14,0.75,1.75,0.08,1.14,-0.36,-0.63,-0.66,-0.14,-1.23,0.48,1.44,0.07,-0.52,-0.76,-0.82,1.06,-1.65,-1.05,0.72,1.56,-1.12,1.7,0.25,-1.22,2.18,0.32,-1.45,0.75,0.76,-0.28,-0.01,-0.61,0.26,-0.85,-1.75,0.24,-1.79,-1.79,-0.65,-0.21,-1.54,0.73,0.83,0.83,-0.76,1.63,1.32,0.05,-1.66,-1.12,1.7,-0.57,1.22,0.4,0.54,-0.55,-1.71,-1,-0.11,0.92,-0.39,0.21,2.56,1.33,-1.82,-0.21,3.19,-1.07,-0.54,1.49,-0.08,-1.89,-0.22,0.92,1.59,0.05,2.08,1.3,0.69,1.21,0.11,2.91,-1.01,-2.62,-0.05,0.18,-1.69,-1,0.71,-0.28,-0.87,0.82,0.08,2.48,-1.14,0.56,-0.34,0.23,0.51,1.64,0.49,-0.15,-1.89,1.22,-0.75,-0.97,0.48,-0.13,-0.05,-1.65,0.27,-0.45,1.46,0.28,-0.15,-0.57,0.44,0.32,-1.06,-0.14,-0.8,-0.3,-0.09,-2.45,1.88,-0.01,-0.36,0.11,0.13,-0.93,-0.82,1.54,-1.23,-1.72,-1.16,0.38,0.28,-2.11,0.51,0.03,1.7,1.71,0.87,-0.83,-0.21,-2.16,-0.47,0.95,-0.68,0.31,0.41,-1.58,0.6,-2.11,-1.48,0.86,0.58,0.26,-0.34,2.56,-0.94,-0.48,0.65,-0.91,0.24,1.06,1.49,-0.12,0.17,-2.43,-2.39,-0.56,-1.32,0.41,-1.01,-0.57,0.76,0.1,-1.48,2.48,0.26,-0.36,1.07,-0.34,-1.03,-1.25,1.24,0.21,-0.25,-0.72,-0.18,1.11,0.59,1.31,1.09,-1.24,0.59,0.83,-0.97,-2.39,0.37,0.72,0.11,1.22,-1.28,-0.51,-0.08,-1.69,-1.71,-1.02,-0.2,-0.18,0.2,-0.96,0.77,-0.14,-0.62,0.57,-1.7,-0.48,-0.56,0.71,0.19,0.1,-0.16,0.05,2.08,-0.11,1.74,0.27,-0.89,-1.56,0.79,-0.73,-0.21,0.18,0.09,-1.04,-1.17,-1.06,-0.03,-1.75,-0.99,0.54,-1.56,-0.72,-0.86,2.09,-0.01,1.09,-0.12,0.02,-0.19,0.63,0.51,0.05,1.26,-0.77,-0.27,1.66,-0.9,-0.66,1.7,-0.89,0.18,1.32,-2.16,-0.61,1.25,0.25,0.71,0.82,2.48,1.44,-1.95,-0.16,0.23,0.46,-0.31,2.4,0.11,1.34,0.19,0.24,0.49,-1.89,-1.23,-0.84,-0.86,-1.46,0.6,1.56,-1.27,0.21,-0.31,-0.18,0.31,0.01,-0.52,0.59,-1.25,-1.75,-0.62,0.04,-0.48,-0.56,1.31,-1.48,-0.06,0.74,-1.06,0.22,-0.69,0.22,-0.19,0.69,-0.52,0.63,-0.58,-1.32,-1.28,-0.27,-0.99,-0.21,-0.05,0.29,1.33,0.95,0,-0.09,-0.28,0.17,-0.62,1.42,1.24,0.46,-0.9,0.21,0.13,-0.2,0.24,0.08,1.44,-1.89,1.57,1.56,1.12,0.63,0.37,0.19,0.84,2.91,1.34,1.17,-1.12,0.95,0.46,-0.08,-0.64,-0.8,1.83,0.2,-1.58,0.58,0.2,-1.23,2.25,-0.02,1.31,1.97,0.17,0.11,0.84,-0.26,-1.46,0.37,-0.3,0.05,0.31,0.19,0.86,-1.12,-0.62,0.79,-0.05,-3.16,-4.24,-1.83,-0.62,-1.06,1.11,-0.08,-1.58,-0.39,0.86,-0.32,-0.58,-0.28,-1.82,-2.27,2.02,0.63,2.25,-1.07,-0.2,-1.3,-0.44,1.84,0.27,-1.12,-0.39,-0.47,-1.28,-0.35,-0.71,0.32,1.05,-1.19,-0.56,-0.88,-1.76,-1.14,-0.05,-0.13,-1.18,1.24,1.43,0.59,-0.65,2.02,-1.48,0.87,-0.54,0.09,0.91,-0.44,-0.53,1.35,1.84,-0.32,-0.9,-0.3,-0.87,1.6,0.4,1.05,-1.16,1.87,0.25,0.51,0.03,0.53,-0.85,0.69,0.92,-1.53,-0.89,0.59,-0.53,-1.68,0.49,0.98,0.22,1.61,1.38,1.59,-0.33,0.79,1.35,-1.19,-1.27,-1.28,1.88,2.91,-0.29,0.88,0.29,0.83,-1.69,0.51,-0.8,-0.68,-0.75,-1.45,-1.37,-0.06,0.55,-0.77,-0.29,1.64,-0.79,-0.63,1.07,0.1,-0.02,-0.52,-1.21,0.28,1.56,0.4,0.43,-0.27,0.36,-0.73,-0.21),
							area_wealth = c(-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1))
```

```{r question_4, exercise=TRUE,exercise.lines = 6}
#language.data

formula <- '

'
```

```{r question_4-hint-1}
Add random slopes within the brackets on the left hand side of the `|` symbol.
```

```{r question_4-hint-2}
Number of siblings does not vary within participants, therefore, you do not have to specify a random slope of siblings in participant_id.

Next hint will show the solution...
```

```{r question_4-solution}
formula <- 'language ~ family_ses + area_wealth +
  (1 + family_ses | school_id)'
```

```{r question_4-code-check}
grade_code()
```

### Question 5

Using the `lmer()` function, run the linear mixed effect model using the formula and data (`language.data`) from Question 4, then answer the following question:

```{r question_5-setup}
language.data <- data.frame(school_id = c(1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,1,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,2,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,3,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,4,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,5,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,6,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,7,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,8,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,9,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,10,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,11,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,12,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,13,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,14,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,15,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,16,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,17,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,18,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,19,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,20,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,21,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,22,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,23,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,24,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,25,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,26,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,27,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,28,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,29,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30,30),
							child_id = c(11,13,16,51,106,145,161,182,291,297,305,318,340,353,415,464,479,494,521,560,578,623,665,704,708,709,800,836,843,894,3,9,30,37,40,64,135,141,143,163,184,220,246,317,330,331,360,380,391,433,446,452,462,471,550,729,765,771,792,884,16,39,57,61,75,99,143,161,165,211,255,281,285,338,344,351,354,359,392,469,473,488,585,734,782,812,828,837,855,896,6,14,42,53,65,85,130,141,174,178,183,198,220,242,282,293,301,326,343,345,593,594,595,603,605,625,743,746,832,853,36,47,84,86,112,113,131,162,228,293,330,358,434,449,478,505,507,509,542,597,602,603,661,676,711,732,737,758,802,840,2,26,49,54,69,81,98,139,150,156,170,289,392,466,486,495,513,532,565,582,590,658,661,674,687,769,770,780,864,892,2,61,177,183,201,226,272,365,423,441,445,482,492,512,523,533,548,582,594,664,737,765,768,810,842,849,850,855,863,873,67,78,88,101,108,169,172,179,184,232,237,239,270,355,368,390,496,542,565,591,594,647,651,659,664,697,719,830,864,877,17,27,43,72,88,114,126,251,275,291,307,323,350,370,399,421,545,572,573,575,590,595,612,654,657,761,792,831,863,886,29,36,144,198,210,221,250,253,316,362,374,394,429,475,482,511,542,577,580,584,594,627,637,685,694,724,779,818,820,846,13,44,66,70,72,73,125,135,194,272,352,464,468,479,492,506,546,566,573,587,603,610,619,626,645,686,714,781,834,865,3,37,41,68,83,191,215,221,238,256,269,278,282,306,324,365,387,405,424,465,525,530,544,560,602,604,628,646,801,868,5,26,30,61,64,119,134,168,185,301,328,334,378,443,450,475,536,542,551,583,608,640,655,703,731,754,768,775,813,876,7,64,92,100,104,123,149,155,172,246,256,293,317,320,442,485,519,583,617,639,654,685,726,744,764,771,785,791,792,814,12,89,123,129,140,204,222,233,243,263,298,315,389,414,451,464,466,535,576,617,618,643,742,756,771,785,818,819,867,877,48,85,96,174,184,199,305,318,323,363,374,424,463,474,497,528,581,584,588,605,624,627,672,739,754,798,836,892,898,900,1,77,96,115,140,146,150,201,209,221,223,259,364,377,393,465,467,494,500,513,536,559,608,620,685,741,768,816,826,886,5,11,12,82,104,133,152,153,165,169,204,226,276,306,347,373,483,496,520,560,565,577,702,725,731,785,797,825,865,870,13,15,22,130,137,158,167,193,197,203,224,229,233,259,318,354,408,418,432,473,480,527,538,629,663,693,727,770,888,896,5,129,194,203,209,233,236,257,259,260,278,283,305,357,371,382,397,404,514,525,548,578,616,646,693,710,744,798,819,822,29,60,77,85,143,162,216,264,304,312,349,352,362,390,408,415,434,460,490,493,624,627,647,673,741,803,853,856,862,870,1,17,18,51,63,68,284,316,326,360,371,429,440,503,525,538,544,545,571,593,624,674,690,696,758,761,772,785,803,812,10,13,81,86,88,115,150,209,244,247,267,286,292,301,340,379,448,460,473,491,494,496,553,653,678,706,764,774,835,837,76,137,230,242,250,278,284,352,366,408,415,427,482,498,526,532,542,555,566,567,575,583,632,644,684,760,809,827,870,900,16,33,53,74,228,245,265,272,281,283,292,302,305,306,402,436,444,485,494,501,506,552,571,724,753,775,836,854,879,895,53,57,60,105,106,117,166,167,234,264,320,449,507,514,563,572,610,650,655,678,722,738,742,747,753,757,771,794,856,900,3,42,54,55,63,124,125,167,199,224,237,351,358,363,438,441,446,449,463,554,636,676,701,741,771,776,821,822,855,882,12,45,56,113,118,160,221,239,273,275,283,337,411,414,441,482,486,492,619,635,676,683,695,701,755,761,783,806,849,867,45,47,75,88,91,126,192,243,253,287,309,367,404,447,461,465,487,566,570,597,605,660,675,695,728,733,760,777,836,893,6,16,43,77,91,106,130,151,176,182,196,327,346,356,393,403,409,462,490,507,508,534,577,599,606,671,696,723,862,870),
							language = c(0.19,-1.07,0.43,-0.98,0.26,0.53,0.21,0.36,-0.53,0.21,-1.01,0.55,0.85,0.21,-0.76,-0.74,-0.59,-1.18,-0.36,0.4,-0.55,0.27,1.05,0.45,0.8,0.4,0.67,0.11,-1.03,-0.9,-0.89,0.52,-1.15,0.14,0.62,-0.08,0.45,0.43,-0.4,-0.26,0.51,-0.01,-0.09,0.59,0.6,-0.04,-0.97,0.45,-1.19,-0.53,0.53,0.1,0.52,0.05,0.28,-0.61,0.05,-0.25,-0.63,0.35,0.2,0.4,-0.41,-0.05,0.16,0.32,-0.58,-0.21,-0.29,-0.2,-1.61,-0.13,-0.4,-0.61,0.01,-0.7,0.15,-0.34,-0.84,-0.64,-0.29,0.36,-0.67,-0.76,-0.94,0.22,-0.1,0.61,0.73,-0.57,0.9,0.28,-0.63,-0.21,-0.52,-0.65,0.07,-0.19,-0.25,1.14,0.77,-0.26,-0.11,-0.42,0.77,-0.1,-0.11,0.24,-0.07,0.61,0,-1.39,-0.71,0.45,0.06,-0.11,0.36,0.01,0.72,-0.49,1.25,-0.22,-0.02,0.74,-0.53,1.05,-1.03,0.02,-0.1,-0.31,0.51,0,1.02,0.99,0.66,-0.52,0.06,0.3,-0.52,0.59,-0.81,0.5,-0.11,0.06,0.48,-1.05,0.43,1.5,0.32,-0.41,0.22,1.02,-0.2,-0.98,-0.08,0.48,-0.29,0.44,-0.13,-0.2,0.13,1.2,0.2,0.05,0.7,0.52,-0.15,-0.85,0.7,0.45,0.83,0.3,-1.02,0.08,0.62,-0.06,-0.14,0.08,0.03,-0.45,-0.65,-0.77,-0.67,0.38,-0.53,-0.45,0.06,0.03,-0.15,0.55,0.07,-1.91,-0.27,-1.16,-0.3,-0.5,-1.19,-0.8,-1.02,-0.86,-0.57,-0.41,-0.13,-0.58,-0.35,-0.15,-0.98,-0.72,0.01,-0.86,-0.5,0.56,-0.01,0.05,-0.84,0.87,0.65,0.03,0.11,0.77,-0.09,-0.21,0.06,0.3,0.47,-1.05,-0.7,-0.79,0.48,-0.42,-1.1,-0.17,0.86,-0.66,0.44,0.74,0.75,0.29,0.11,0.74,0.17,-0.42,-0.48,-0.12,-0.08,0.28,0.8,0.13,-0.13,-0.36,0.78,-0.13,0.17,-0.12,-1.87,0.84,0.93,0.1,-0.79,-0.23,0.41,0.04,-0.87,-0.42,-0.25,0.23,-0.41,-0.3,0.12,-0.5,-0.91,0.41,-0.09,-0.48,0.25,-0.03,0.6,-0.09,-0.53,0.1,0.16,0.47,-1.23,-1.15,-1.72,0.45,-0.27,-0.21,0.35,-0.02,-0.3,0.26,0.17,0.21,0.71,-0.18,-0.02,0.35,-0.32,-0.3,-0.95,-0.17,-0.45,-1.54,0.22,0.05,0.58,-0.31,-0.45,-0.35,-0.46,-0.78,-0.36,-0.87,0.49,1.31,-1,0.37,-0.34,-0.26,0.56,-0.74,-0.01,-0.97,0.24,0.49,-0.45,-0.39,-0.09,-1.01,-0.21,0.47,0.27,0.14,0.01,-0.08,0.23,-0.5,0.05,-0.88,0.18,0.22,0.31,-0.79,0.18,1.08,1.14,0.05,0.61,-0.22,-0.35,0.6,-0.82,-0.34,0.22,-1.34,-0.14,0.56,0.18,-0.51,1,-0.45,-0.48,0.19,-0.46,-0.08,0.8,-0.86,0.39,-0.24,0.09,1.17,0.51,-0.58,0.57,-0.33,-0.2,0.5,0.82,0.49,-0.49,-0.43,-0.03,0.2,-0.58,1.82,1.16,0.82,-0.02,0.18,0.76,-0.13,-1.48,0.51,0,0.16,1,-0.27,0.72,-0.31,0.51,0.45,-0.35,-0.62,-0.34,0.44,0.35,-0.21,0.62,-1.24,-0.01,-0.3,-0.28,-0.49,0.85,-0.98,0.82,-0.17,-0.11,0.84,-0.42,-0.74,-0.41,-0.14,0.68,-0.79,-0.65,0.03,-1.2,-0.19,0.7,0.54,0.23,-0.7,-0.44,-0.56,0.16,0.01,-1.19,-1.38,0.8,-0.19,0.58,-0.66,-0.6,0.42,-0.52,-0.14,0.37,0.55,0.13,-0.04,0.09,0.13,0.16,0.74,-0.23,0.64,1.04,-0.01,0.65,0.69,0.02,0.24,0.56,0.85,0.09,-0.46,1.02,0.94,0.38,-0.4,0.32,0.01,0.64,1.35,-0.07,-0.71,-0.8,0.08,0.11,-0.9,-0.62,0.37,-0.18,0.18,-0.14,-0.74,0.96,0.5,-0.13,-0.7,-0.21,1.9,0.37,0.46,-0.13,-0.3,-0.22,-0.15,0.16,-0.35,0.18,-0.17,-0.57,-0.44,0.74,0.17,0.3,0.04,-0.22,0.32,0.68,0.29,-0.1,-1.11,-0.97,-0.34,-0.7,1.14,-0.89,0.32,-0.75,0.07,-0.51,-0.11,0.02,-1.38,-1.53,-0.72,-0.42,0.31,-0.42,0.21,-0.02,0.38,1.02,0.18,-0.56,-0.83,-1.56,-0.77,0.08,-0.14,-0.81,-0.75,-1.03,0.11,-1.37,-0.84,-0.6,-0.74,-0.48,0.3,0.87,0.22,-0.69,0.33,-0.42,-0.26,-0.53,0.11,-0.95,-1.47,-0.9,-0.6,-0.5,-1.34,1,0.28,-0.33,-0.02,-0.17,-0.74,0.53,-0.32,0.37,0.16,-0.8,0.07,-0.66,0.1,-0.42,-0.7,0.01,-0.26,-0.31,0.5,0.81,0.44,-0.49,-0.11,-0.58,-0.52,-0.92,-0.12,0.44,0.21,0.06,-0.44,-0.72,-0.92,-0.77,-0.71,-0.32,-0.07,-0.43,-0.45,-0.38,0.22,0.11,-0.58,0.44,0.07,-0.49,-0.88,0.54,0.65,-0.46,-0.03,0.48,0.08,-0.11,0.97,0.23,-1.13,-1.2,-0.17,-1.31,-0.52,0.09,0.18,0.09,-0.35,-0.8,1.01,-0.9,-0.66,0.12,-0.76,-0.87,0.99,0.39,-0.34,0.49,-0.19,-0.29,-0.47,0.86,-0.9,-1.18,0.91,0.01,-0.54,1.02,-0.01,-0.22,0.17,-1.3,-0.42,-0.19,-0.46,0.19,1,0.22,0.25,0.12,1.16,0.59,0.11,0.5,1.05,-0.29,0.39,0.38,-0.54,0.71,0.43,-0.07,0.03,0.31,-1.42,0.34,-0.43,-0.09,0.61,0.82,-0.64,0.45,0.08,0.4,0.27,-0.23,-0.23,0.66,-0.43,-1,-0.48,-0.39,0.75,0.54,-0.63,-0.18,-0.85,-0.69,0.35,0.49,0,-0.14,-0.05,-0.16,-0.53,0.24,-0.18,-0.88,-1.09,0.24,-0.42,0.57,-0.52,-0.26,0.38,-0.98,0.19,-0.19,-0.16,0.07,-1.53,-0.11,-0.15,0.42,-0.93,-0.24,0.22,-0.2,0.63,-0.31,0.41,-0.8,0.41,0.57,0.12,-0.32,0.27,0.97,0.58,1.21,0.55,0.65,-0.15,0.98,-0.71,-0.02,-0.53,0.24,0.97,0.14,-0.54,0.78,-0.29,-1.08,0.29,-0.56,0,1.03,-0.12,0.39,0.48,-0.21,-0.9,0.29,0.25,-0.39,0.54,-0.14,-0.21,-1.02,-0.21,0.38,-0.92,-2.46,-1.99,-1.17,-1.15,-1.27,0.72,-1.06,-1,-0.24,0.27,-0.4,-0.81,-0.54,-1.15,-0.88,0.11,-0.36,0.69,-0.73,-1.09,-1.53,-0.72,0.34,-0.05,-0.19,-0.74,-1.29,-0.88,-0.91,-0.64,-0.21,0.38,-0.35,0.06,-0.53,-0.28,-0.26,-0.13,-0.62,-1.25,0.43,0.77,0.1,-0.11,1.41,-0.69,0.34,-1.12,0.3,0.13,0.27,-0.11,0.79,0.64,-0.19,-0.66,-1.15,-0.78,0.78,-0.49,0.07,0.09,0.9,-1.06,0.47,0.72,0.21,0.28,0.48,-0.01,-0.52,0.13,0.46,0.3,0.2,0.47,0.4,-0.34,0.98,-0.06,0.87,-0.11,0.13,0.12,-0.24,-0.3,0.81,0.2,1.35,-0.27,1.09,0.26,0.38,-0.75,0.99,1.24,0.46,-0.33,-0.56,0.23,-0.22,0.1,0.13,0.14,1.38,0.01,0.24,-0.89,0.18,-0.06,-0.03,0.37,0.19,1.07,0.42,0.09,-0.23,0.38,-0.36,-0.08),
							family_ses = c(0.44,-2.16,0.29,-1.17,-0.8,-0.25,0.64,-1.37,-0.35,0.95,0.21,2.56,2.4,0.47,-0.56,-1.54,-1.06,-1.89,0.86,-1.16,0.59,-0.04,1.27,0.95,0.53,0.39,0.36,2.91,0.82,-0.45,-3.16,1.56,-0.54,-0.15,0.12,-0.01,0.91,0.98,-1.02,-1.34,0.92,-0.65,-0.36,-0.14,1.72,0.05,-1.56,-0.1,-0.95,-1.66,0.63,-0.73,1.07,-0.11,1.06,0.96,0.46,-1.12,-1.22,1.16,0.29,-0.93,0.46,-0.03,1.87,0.69,-1.02,0.64,-2.45,0.74,-0.25,1.42,0.29,1.14,-0.22,-0.58,-0.94,-1.07,-1.58,-2.05,0.24,1.25,-1.3,-0.34,-1.82,0.18,0.77,-0.31,-0.35,-1.01,0.88,-0.05,-4.24,0.95,-1.53,-1.71,-0.68,0.98,-0.11,0.79,1.98,0.51,-0.65,-0.52,-0.42,-0.66,-0.31,0.54,1.07,0.64,0.51,-1.68,-1.14,0.36,1.59,0.73,0.81,0.88,0.48,-1.56,1.18,-1.16,0.15,1.25,-0.87,-0.56,-1.45,-0.2,-0.09,-0.66,1.72,-0.28,0.71,2.25,0.85,-1.34,-0.02,0.47,0.22,1.38,-0.81,0.36,0.19,-0.44,0.81,-0.73,0.26,1.66,0.89,-1.5,0.17,-0.42,-0.91,-1.83,0.17,-0.61,1.35,0.48,0.82,-0.41,0.2,0.7,-1.58,0.73,0.87,0.98,-0.75,-1.06,0.38,0.44,1.65,0.39,0.19,1.26,0.09,-0.77,-1.32,-2.62,0.17,-1.01,0.17,-0.03,-0.73,1.98,0.08,-0.36,-0.62,2.07,0.69,2.02,-0.02,-1.48,-0.54,-1.42,-0.49,-0.02,-1.24,0.44,-1.68,0.52,0.26,0.46,-0.45,1.75,0.38,1.6,-0.66,-0.35,0.89,-0.06,0.26,0.3,0.25,1.25,-1.49,1.88,1.14,1.57,0.92,1.11,-0.32,-0.05,0.72,2.17,0.75,-1.7,-1.23,0.22,0.38,0.6,-1.68,-0.11,-0.8,-0.63,0.52,1.48,0.47,0.05,0.17,0.54,0.09,-0.36,0.83,-0.36,0.25,0.37,0.03,0.34,-1.18,-0.35,-0.71,1.33,0.36,-1.14,-1.93,1.71,-0.19,0.17,-0.39,0.69,1.65,-1.14,0.67,1.06,-1.87,-0.9,-1.22,-1.31,0.89,-0.15,-0.51,1.18,-0.38,0.51,-0.42,-1.14,0.59,0.69,-0.99,0.57,-0.21,2.65,-0.86,-1.46,-1.48,0.38,0.22,0.28,0.92,-0.22,-1.68,2.08,-0.88,-1.65,0.08,0.37,0.76,-0.57,-0.49,-0.78,-2.16,0.06,-0.38,-0.68,-0.36,1.14,-0.08,0.91,0.1,-0.62,-0.62,-1.54,0.48,-1.06,-0.54,1.56,0.05,0.22,-0.39,-0.53,0.36,0.11,0.09,-0.17,0.74,0.81,1.75,1.26,1.44,-0.83,-3.16,-0.15,1.1,-0.03,1.07,-1.6,0.04,-1.14,-0.12,-0.63,-0.28,-1.25,-0.42,0.13,-0.53,2.07,0.78,0.52,3.19,0.49,1.09,0.87,0.02,-1.16,-0.81,-0.95,-0.18,-0.97,0.41,-0.4,-0.57,-0.42,-0.54,-0.03,-0.01,-1.08,1.39,0.54,0.65,-0.31,-0.41,0.17,-1.35,-0.05,0.57,-1.46,-0.97,0.22,0.59,-0.52,-0.13,-1.22,-0.26,-0.17,0.03,1.21,-0.45,0.84,-1.04,0.97,1.08,-0.01,-1.67,0.17,-0.14,0.75,1.75,0.08,1.14,-0.36,-0.63,-0.66,-0.14,-1.23,0.48,1.44,0.07,-0.52,-0.76,-0.82,1.06,-1.65,-1.05,0.72,1.56,-1.12,1.7,0.25,-1.22,2.18,0.32,-1.45,0.75,0.76,-0.28,-0.01,-0.61,0.26,-0.85,-1.75,0.24,-1.79,-1.79,-0.65,-0.21,-1.54,0.73,0.83,0.83,-0.76,1.63,1.32,0.05,-1.66,-1.12,1.7,-0.57,1.22,0.4,0.54,-0.55,-1.71,-1,-0.11,0.92,-0.39,0.21,2.56,1.33,-1.82,-0.21,3.19,-1.07,-0.54,1.49,-0.08,-1.89,-0.22,0.92,1.59,0.05,2.08,1.3,0.69,1.21,0.11,2.91,-1.01,-2.62,-0.05,0.18,-1.69,-1,0.71,-0.28,-0.87,0.82,0.08,2.48,-1.14,0.56,-0.34,0.23,0.51,1.64,0.49,-0.15,-1.89,1.22,-0.75,-0.97,0.48,-0.13,-0.05,-1.65,0.27,-0.45,1.46,0.28,-0.15,-0.57,0.44,0.32,-1.06,-0.14,-0.8,-0.3,-0.09,-2.45,1.88,-0.01,-0.36,0.11,0.13,-0.93,-0.82,1.54,-1.23,-1.72,-1.16,0.38,0.28,-2.11,0.51,0.03,1.7,1.71,0.87,-0.83,-0.21,-2.16,-0.47,0.95,-0.68,0.31,0.41,-1.58,0.6,-2.11,-1.48,0.86,0.58,0.26,-0.34,2.56,-0.94,-0.48,0.65,-0.91,0.24,1.06,1.49,-0.12,0.17,-2.43,-2.39,-0.56,-1.32,0.41,-1.01,-0.57,0.76,0.1,-1.48,2.48,0.26,-0.36,1.07,-0.34,-1.03,-1.25,1.24,0.21,-0.25,-0.72,-0.18,1.11,0.59,1.31,1.09,-1.24,0.59,0.83,-0.97,-2.39,0.37,0.72,0.11,1.22,-1.28,-0.51,-0.08,-1.69,-1.71,-1.02,-0.2,-0.18,0.2,-0.96,0.77,-0.14,-0.62,0.57,-1.7,-0.48,-0.56,0.71,0.19,0.1,-0.16,0.05,2.08,-0.11,1.74,0.27,-0.89,-1.56,0.79,-0.73,-0.21,0.18,0.09,-1.04,-1.17,-1.06,-0.03,-1.75,-0.99,0.54,-1.56,-0.72,-0.86,2.09,-0.01,1.09,-0.12,0.02,-0.19,0.63,0.51,0.05,1.26,-0.77,-0.27,1.66,-0.9,-0.66,1.7,-0.89,0.18,1.32,-2.16,-0.61,1.25,0.25,0.71,0.82,2.48,1.44,-1.95,-0.16,0.23,0.46,-0.31,2.4,0.11,1.34,0.19,0.24,0.49,-1.89,-1.23,-0.84,-0.86,-1.46,0.6,1.56,-1.27,0.21,-0.31,-0.18,0.31,0.01,-0.52,0.59,-1.25,-1.75,-0.62,0.04,-0.48,-0.56,1.31,-1.48,-0.06,0.74,-1.06,0.22,-0.69,0.22,-0.19,0.69,-0.52,0.63,-0.58,-1.32,-1.28,-0.27,-0.99,-0.21,-0.05,0.29,1.33,0.95,0,-0.09,-0.28,0.17,-0.62,1.42,1.24,0.46,-0.9,0.21,0.13,-0.2,0.24,0.08,1.44,-1.89,1.57,1.56,1.12,0.63,0.37,0.19,0.84,2.91,1.34,1.17,-1.12,0.95,0.46,-0.08,-0.64,-0.8,1.83,0.2,-1.58,0.58,0.2,-1.23,2.25,-0.02,1.31,1.97,0.17,0.11,0.84,-0.26,-1.46,0.37,-0.3,0.05,0.31,0.19,0.86,-1.12,-0.62,0.79,-0.05,-3.16,-4.24,-1.83,-0.62,-1.06,1.11,-0.08,-1.58,-0.39,0.86,-0.32,-0.58,-0.28,-1.82,-2.27,2.02,0.63,2.25,-1.07,-0.2,-1.3,-0.44,1.84,0.27,-1.12,-0.39,-0.47,-1.28,-0.35,-0.71,0.32,1.05,-1.19,-0.56,-0.88,-1.76,-1.14,-0.05,-0.13,-1.18,1.24,1.43,0.59,-0.65,2.02,-1.48,0.87,-0.54,0.09,0.91,-0.44,-0.53,1.35,1.84,-0.32,-0.9,-0.3,-0.87,1.6,0.4,1.05,-1.16,1.87,0.25,0.51,0.03,0.53,-0.85,0.69,0.92,-1.53,-0.89,0.59,-0.53,-1.68,0.49,0.98,0.22,1.61,1.38,1.59,-0.33,0.79,1.35,-1.19,-1.27,-1.28,1.88,2.91,-0.29,0.88,0.29,0.83,-1.69,0.51,-0.8,-0.68,-0.75,-1.45,-1.37,-0.06,0.55,-0.77,-0.29,1.64,-0.79,-0.63,1.07,0.1,-0.02,-0.52,-1.21,0.28,1.56,0.4,0.43,-0.27,0.36,-0.73,-0.21),
							area_wealth = c(-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,-0.64,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-1.37,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,-0.19,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.71,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,0.41,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-2.07,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.31,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.22,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-0.32,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,-1.55,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.25,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.35,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,0.07,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,0.1,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-0.53,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.61,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-1.58,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.17,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,-0.75,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,0.42,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.6,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.46,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-0.43,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,-2.36,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.04,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,0.13,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1,1.1))
```

```{r question_5, exercise=TRUE,exercise.lines = 12}
#language.data

formula <- '

'

lmer() %>%
  summary()

```

```{r question_5-hint-1}
Use the formula from Question 4 in the `lmer()` function.
```

```{r question_5-hint-2}
Make sure to read the results using the `summary()` function.

Next hint will show the solution...
```

```{r question_5-solution}
formula <- 'language ~ family_ses + area_wealth +
  (1 + family_ses | school_id)'
  
lmer(formula,data = language.data) %>%
  summary()
```

```{r question_5-quiz}
quiz(question("Select all the correct interpretations about the model above:",
  answer("There is a significant positive effect of family SES on language development.",correct = TRUE),
  answer("There is a significant positive effect of area wealth of school on language development.",correct = TRUE),
  answer("There is no significant effect of family SES on language development"),
  answer("There is no significant effect of area wealth of school on language development."),allow_retry = TRUE,random_answer_order = TRUE))
```

### Question 6

You conduct a study investigating threat perceptions of faces. 40 participants were asked to rate 30 faces. You measure the facial masculinity of the 30 faces, and participants reported their own locus of control (higher scores indicating greater external locus of control). You predict that more masculine faces will be perceived as threatening, but more so for participants with greater external locus of control.

You collect your data, and it is currently saved in short-form as `threat.data`. Reshape the data so that it is in long form and save it as the object `threat.data_long`. Name the new key variable "face_id" and the new value variable "threat".

```{r}
threat.data
```

```{r question_6-setup}

```

```{r question_6, exercise=TRUE,exercise.lines = 5}
#threat.data

threat.data_long <- 
```

```{r question_6-hint-1}
Use the `gather()` function to convert data from short-form to long-form.
```

```{r question_6-hint-2}
If you are struggling to use the `gather()` function, try googling examples that you can copy and edit.

Next hint will show the solution...
```

```{r question_6-solution}
threat.data_long <- gather(threat.data,key = "face_id",value = "threat",face_1:face_30)

#Note: "face_1:face_30" is shorthand of saying all variables between 'face_1' and 'face_30'.
```

```{r question_6-code-check}
grade_code()
```

### Question 7

The masculinity of each face from the study in Question 6 is saved as the data.frame `face.data` (shown below). Using the `left_join()` function, combine this data.frame with the long-form data from Question 6. Save the result as `analysis.data`.

```{r}
face.data <- data.frame(face_id = paste("face_",1:30,sep = ""),
	masc = c(-1.25,0.74,-0.08,0.21,0.4,-0.09,-0.55,-2.22,1.47,-1.75,-0.51,-0.26,-0.51,-1.58,0.56,-0.09,1.17,0.08,0.63,0.89,-2.01,1.02,-0.56,0.06,0.2,0.83,0.54,0.44,0.6,1.41),stringsAsFactors = FALSE)
	
face.data
```

```{r question_7-setup}
threat.data_long <- gather(threat.data,key = "face_id",value = "threat",face_1:face_30)

face.data <- data.frame(face_id = paste("face_",1:30,sep = ""),
	masc = c(-1.25,0.74,-0.08,0.21,0.4,-0.09,-0.55,-2.22,1.47,-1.75,-0.51,-0.26,-0.51,-1.58,0.56,-0.09,1.17,0.08,0.63,0.89,-2.01,1.02,-0.56,0.06,0.2,0.83,0.54,0.44,0.6,1.41),stringsAsFactors = FALSE)
```

```{r question_7, exercise=TRUE,exercise.lines = 5}
#threat.data_long
#face.data

analysis.data <- 
```

```{r question_7-hint-1}
The `left_join()` function. is used to join to data.frames. Either use Google or type help(left_join) into the console see how to use this function.
```

```{r question_7-hint-2}
You will want to put the `threat.data_long` data on the left-hand side, and the `face.data` data on the right-hand side. The data is joined by the variable 'face_id'.

Next hint will show the solution...
```

```{r question_7-solution}
analysis.data <- left_join(threat.data_long,face.data,by = "face_id")
```

```{r question_7-code-check}
grade_code()
```

### Question 8

Using the `analysis.data` data.frame created in Question 7, create an lmer formula for the analysis to test for the interaction between facial masculinity and rater locus of control on threat judgements. Note that variables have already been standardised at the appropriate level, so they do not need to be centered. Specify random slopes maximally.

```{r}
threat.data_long <- gather(threat.data,key = "face_id",value = "threat",face_1:face_30)

face.data <- data.frame(face_id = paste("face_",1:30,sep = ""),
	masc = c(-1.25,0.74,-0.08,0.21,0.4,-0.09,-0.55,-2.22,1.47,-1.75,-0.51,-0.26,-0.51,-1.58,0.56,-0.09,1.17,0.08,0.63,0.89,-2.01,1.02,-0.56,0.06,0.2,0.83,0.54,0.44,0.6,1.41),stringsAsFactors = FALSE)
	
analysis.data <- left_join(threat.data_long,face.data,by = "face_id")

analysis.data
```

```{r question_8-setup}
threat.data_long <- gather(threat.data,key = "face_id",value = "threat",face_1:face_30)

face.data <- data.frame(face_id = paste("face_",1:30,sep = ""),
	masc = c(-1.25,0.74,-0.08,0.21,0.4,-0.09,-0.55,-2.22,1.47,-1.75,-0.51,-0.26,-0.51,-1.58,0.56,-0.09,1.17,0.08,0.63,0.89,-2.01,1.02,-0.56,0.06,0.2,0.83,0.54,0.44,0.6,1.41),stringsAsFactors = FALSE)
	
analysis.data <- left_join(threat.data_long,face.data,by = "face_id")
```

```{r question_8, exercise=TRUE,exercise.lines = 5}
#analysis.data

formula <- '

'
```

```{r question_8-hint-1}
Use the `*` symbol to specify an interaction, exactly as you would for a normal linear regression.
```

```{r question_8-hint-2}
When specifying the random slopes, 'masc' does not vary within 'face_id' and 'locus_of_control' does not vary within 'participant_id'. Therefore, these do not need to be specified, even as part of an interaction.

Next hint will show the solution...
```

```{r question_8-solution}
formula <- '
  threat ~ masc*locus_of_control +
    (1 + masc | participant_id) +
    (1 + locus_of_control | face_id)
'
```

```{r question_8-code-check}
grade_code()
```

### Question 9

Using the `lmer()` function, run the linear mixed effect model using the formula and data from Question 8, then answer the following question:

```{r question_9-setup}
threat.data_long <- gather(threat.data,key = "face_id",value = "threat",face_1:face_30)

face.data <- data.frame(face_id = paste("face_",1:30,sep = ""),
	masc = c(-1.25,0.74,-0.08,0.21,0.4,-0.09,-0.55,-2.22,1.47,-1.75,-0.51,-0.26,-0.51,-1.58,0.56,-0.09,1.17,0.08,0.63,0.89,-2.01,1.02,-0.56,0.06,0.2,0.83,0.54,0.44,0.6,1.41),stringsAsFactors = FALSE)
	
analysis.data <- left_join(threat.data_long,face.data,by = "face_id")
```

```{r question_9, exercise=TRUE,exercise.lines = 12}
#analysis.data

formula <- '
  threat ~ masc*locus_of_control +
    (1 + masc | participant_id) +
    (1 + locus_of_control | face_id)
'

lmer() %>%
  summary()

```

```{r question_9-hint-1}
Use the formula from Question 8 in the `lmer()` function.
```

```{r question_9-hint-2}
Make sure to read the results using the `summary()` function.

Next hint will show the solution...
```

```{r question_9-solution}
formula <- '
  threat ~ masc*locus_of_control +
    (1 + masc | participant_id) +
    (1 + locus_of_control | face_id)
'

model <- lmer(formula,data = analysis.data)
summary(model)
```

```{r question_9-code-check}
grade_code()
```

```{r question_9-quiz}
quiz(question("The hypothesis that more masculine faces will be perceived as more threatening, but more so for participants with greater external locus of control, was...",
  answer("supported.",correct = TRUE),
  answer("not supported."),allow_retry = TRUE,random_answer_order = TRUE))
```



